```python
#Initialize Otter

import otter
grader = otter.Notebook("DSA2024Test.ipynb")
```

## DSA 2024 Summer School Admittance Check

Thanks for your interest in attending DSA 2024 Nyeri, Kenya. To attend the summer school you have to have some level of basic Python proficiency. Completing the following notebook should ensure you have the right kind of background to benefit maximally from the Summer School. See you in Nyeri!

```python
import pandas as pd
import numpy as np
%matplotlib inline
import otter

grader = otter.Notebook()
```

**Question 1:** write a function `isValid(s)` that takes as argument a string s containing a sequence of parenthesis '(', ')', '{', '}', '[' and ']', and  determines if the input is valid. A input string is valid if for every open parenthensis there is a close one and parenthesis is well-formed. e.g  "(){}[]" is valid.

```python
def isValid(s: str) -> bool:
    stack = []
    mapping = {')': '(', '}': '{', ']': '['}
    
    for char in s:
        if char in mapping:
            top_element = stack.pop() if stack else '#'
            if mapping[char] != top_element:
                return False
        else:
            stack.append(char)
    
    return not stack

# Test the function
print(isValid("(){}[]"))  # Output should be True

```

Output:
```
True
```

```python
grader.check("q1")
```

Output:
```
q1 results: All test cases passed!
```

**Question 2:** Given a paragraph as a string, write a function that return the number of character with odd frequencies. E.g The paragraph ``DSA 2024 Nyeri`` has *10* characters with odd frequencies. i.e the entire frequency count is given as {' ': 2, '2': 2, 'D': 1, 'S': 1, 'A': 1, '0': 1, '4': 1, 'N': 1, 'y': 1, 'e': 1, 'r': 1, 'i': 1}) and there are *10* characters with odd frequences. So the function should return *10*. 

```python
def oddFrequencyCounter(theParagraph):
    char_count = {}
    for char in theParagraph:
        char_count[char] = char_count.get(char, 0) + 1
    odd_count = sum(1 for count in char_count.values() if count % 2 != 0)
    return odd_count

# Test the function
print(oddFrequencyCounter("DSA 2024 Nyeri"))  # 10

```

Output:
```
10
```

```python
grader.check("q2")
```

Output:
```
q2 results: All test cases passed!
```

**Question 3:** Write an infinite generator function `odd_squares_sum` that yields the sum of square of odd numbers. e.g $1^2 + 3^2 + 5^2 + ...$ up to a ``limit``

```python
def odd_squares_sum(limit):
    total = 0
    num = 1
    while True:
        total += num ** 2
        yield num ** 2
        num += 2
        if num >limit:
            break

# Testing
limit=100
sum_generator = odd_squares_sum(limit)
for value in sum_generator:
    print(value)


```

Output:
```
1
9
25
49
81
121
169
225
289
361
441
529
625
729
841
961
1089
1225
1369
1521
1681
1849
2025
2209
2401
2601
2809
3025
3249
3481
3721
3969
4225
4489
4761
5041
5329
5625
5929
6241
6561
6889
7225
7569
7921
8281
8649
9025
9409
9801
```

```python
grader.check("q3")
```

Output:
```
q3 results: All test cases passed!
```

**Question 4:** Using the `odd_squares_sum` generator defined above, create a list of sum of squares up to a limit of $20$ and store the results in a numpy.array variable called `oddSumList` 

```python
limit = 20
#The np.fromiter function takes an iterable (in this case, the generator odd_squares_sum(limit)) and a data type (dtype=int) and creates a NumPy array from the yielded values.
oddSumList = np.fromiter(odd_squares_sum(limit), dtype=int)

# Testing
print(oddSumList)


```

Output:
```
[  1   9  25  49  81 121 169 225 289 361]
```

```python
grader.check("q4")
```

Output:
```
q4 results: All test cases passed!
```

**Question 5:** Compute the element-wise remainder of ``oddSumList`` when divided by $5$ and merge it with ``oddSumList``. The final output stored in the variable `mergedList` should be in the form of a list of tupples e.g ``[(1,1), (4,9), (0,25), ...]`` 

```python
remainder_list = oddSumList % 5
mergedList = list(zip(remainder_list, oddSumList))

# Testing
print(mergedList)




```

Output:
```
[(1, 1), (4, 9), (0, 25), (4, 49), (1, 81), (1, 121), (4, 169), (0, 225), (4, 289), (1, 361)]
```

```python
grader.check("q5")
```

Output:
```
q5 results: All test cases passed!
```

**Question 6:**  Write a function `greatest_common_divisor` that takes two inputs `a` and `b` and returns the greatest common divisor of the two numbers. E.g. input `(10, 15)` would return `5`

```python
def greatest_common_divisor(a, b):
    while b:
        a, b = b, a % b
    return a

# Test the function
print(greatest_common_divisor(10, 15))  # Output: 5
print(greatest_common_divisor(15, 20))

```

Output:
```
5
5
```

```python
grader.check("q6")
```

Output:
```
q6 results: All test cases passed!
```

**Question 7:**  Write a function `get_3_nearest` that takes in a point of interest ``pt`` and a **list** of points ``ptlist``  and returns a list of 3 nearest points from the point of interest ``pt``. Assume the distance between any two point is defined by the `L1-norm`.

```python
#We'll calculate the Manhattan distance (L1-norm) between the point of interest and each point in the list.
#Then, we'll sort the points based on their distances and return the 3 nearest ones.
def get_3_nearest(pt, ptlist):
    sorted_points = sorted(ptlist, key=lambda x: abs(x[0] - pt[0]) + abs(x[1] - pt[1]))
    return sorted_points[:3]

# Testing
print(get_3_nearest((0, 0), [(1, 1), (2, 2), (-1, -1), (3, 3)]))


```

Output:
```
[(1, 1), (-1, -1), (2, 2)]
```

```python
grader.check("q7")
```

Output:
```
q7 results: All test cases passed!
```

**Question 8:**  Write a function `diagonal_vector(M)` that returns a **numpy** array of the list of **absolute** values of the main diagonal entries in the matrix $M$

```python
def diagonal_vector(M):
    return np.abs(np.diag(M))

# Test the function
M = np.array([[1, 2, 3],
              [4, 5, 6],
              [7, 8, 9]])
print(diagonal_vector(M))  # Output: [1 5 9]

```

Output:
```
[1 5 9]
```

```python
grader.check("q8")
```

Output:
```
q8 results: All test cases passed!
```

**Question 9:**  Write a function `flatten_reverse_lists` that takes in a list of lists and outputs a **reverse** sorted list of elements of sublists of the input list (confusing right?) <br>
Example: given `flatten_reverse_lists([[2,13,44], [6,7]])` it should return `[2,6,7,13,44]`

```python
def flatten_reverse_lists(superlist):
    flattened_list = [element for sublist in superlist for element in sublist]
    return sorted(flattened_list, reverse=True)

# Test the function
print(flatten_reverse_lists([[2, 13, 44], [6, 7]]))  # Output: [44, 13, 7, 6, 2]

```

Output:
```
[44, 13, 7, 6, 2]
```

```python
grader.check("q9")
```

Output:
```
q9 results: All test cases passed!
```

**Question 9:** Create a DataFrame mirroring the table below and assign this to `data`. 

| flavor | scoops | price |
|-----|-----|-----|
| white chocolate | 1 | 2 |
| vanilla | 1 | 1.5 |
| dark chocolate | 2 | 3 |
| strawberry | 1 | 2 |
| strawberry | 3 | 4 |
| vanilla | 2 | 2 |
| mint | 1 | 4 |
| mint | 2 | 5 |
| white chocolate | 3 | 2 |
| dark chocolate | 3 | 3 |
| white chocolate | 2 | 2 |
| dark chocolate | 5 | 3 |


```python
import pandas as pd

data = pd.DataFrame({
    'flavor': ['white chocolate', 'vanilla', 'dark chocolate', 'strawberry', 'strawberry', 
               'vanilla', 'mint', 'mint', 'white chocolate', 'dark chocolate', 
               'white chocolate', 'dark chocolate'],
    'scoops': [1, 1, 2, 1, 3, 2, 1, 2, 3, 3, 2, 5],
    'price': [2, 1.5, 3, 2, 4, 2, 4, 5, 2, 3, 2, 3]
})
df=pd.DataFrame(data)
print(df)
```

Output:
```
flavor  scoops  price
0   white chocolate       1    2.0
1           vanilla       1    1.5
2    dark chocolate       2    3.0
3        strawberry       1    2.0
4        strawberry       3    4.0
5           vanilla       2    2.0
6              mint       1    4.0
7              mint       2    5.0
8   white chocolate       3    2.0
9    dark chocolate       3    3.0
10  white chocolate       2    2.0
11   dark chocolate       5    3.0
```

**Question 10:** Do the following to the dataframe:
* Create a new collumn ``total_price`` whose value is equal to $scoops * price$*
* Write a function ``groupStatistics(data, groupValue)``. Internally, this function groups ``data``  by ``flavor`` and then returns statistics of a given grouped item ``groupValue`` indexed on the ``total_price`` columns. The statistics is a numpy array contains ``[mean, media, min, max, std]`` of the ``total_price`` column. The ``std`` should be rounded to 2 **decimal places**



```python
import pandas as pd

data = pd.DataFrame({
    'flavor': ['white chocolate', 'vanilla', 'dark chocolate', 'strawberry', 'strawberry', 
               'vanilla', 'mint', 'mint', 'white chocolate', 'dark chocolate', 
               'white chocolate', 'dark chocolate'],
    'scoops': [1, 1, 2, 1, 3, 2, 1, 2, 3, 3, 2, 5],
    'price': [2, 1.5, 3, 2, 4, 2, 4, 5, 2, 3, 2, 3]
})

# Adding total_price column
data['total_price'] = data['scoops'] * data['price']

# Function to calculate group statistics
def groupStatistics(data, groupValue):
    grouped_data = data.groupby('flavor').get_group(groupValue)
    mean = grouped_data['total_price'].mean()
    median = grouped_data['total_price'].median()
    min_price = grouped_data['total_price'].min()
    std = round(grouped_data['total_price'].std(),2)
    return np.array([mean, median, min_price, std])

# Testing
print(groupStatistics(data, 'strawberry'))
print(groupStatistics(data, 'mint'))
```

Output:
```
[7.   7.   2.   7.07]
[7.   7.   4.   4.24]
```

```python
grader.check("q11")
```

Output:
```
q11 results: All test cases passed!
```

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

Download the exported ZIP. Take note of the ZIP number and proceed to fill the summer school form

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(pdf=False)
```

Output:
```
<IPython.core.display.HTML object>
```

Tests for q1 (points: 4)
```
>>> points: 1
>>> 
>>> def test_isValid(isValid):
...     assert isValid('{[()]}}') == False
>>> test_isValid(isValid)
```
```
>>> points: 1
>>> 
>>> def test_isValid(isValid):
...     assert isValid('') == True
>>> test_isValid(isValid)
```
```
>>> points: 1
>>> 
>>> def test_isValid(isValid):
...     assert isValid('(){}[]') == True
>>> test_isValid(isValid)
```


Tests for q11 (points: 5)
```
>>> points: 1
>>> np.testing.assert_array_equal(groupStatistics(data, 'strawberry'), [7.0, 7.0, 2.0, 7.07])
```
```
>>> points: 1
>>> np.testing.assert_array_equal(groupStatistics(data, 'dark chocolate'), [10.0, 9.0, 6.0, 4.58])
```
```
>>> points: 1
>>> np.testing.assert_array_equal(groupStatistics(data, 'white chocolate'), [4.0, 4.0, 2.0, 2.0])
```
```
>>> points: 1
>>> np.testing.assert_array_equal(groupStatistics(data, 'vanilla'), [2.75, 2.75, 1.5, 1.77])
```
```
>>> points: 1
>>> np.testing.assert_array_equal(groupStatistics(data, 'mint'), [7.0, 7.0, 4.0, 4.24])
```


Tests for q2 (points: 4)
```
>>> points: 1
>>> 
>>> def test_oddFrequencyCounter(oddFrequencyCounter):
...     paragraph = 'DSA 2024 Nyeri'
...     assert oddFrequencyCounter(paragraph) == 10
```
```
>>> points: 1
>>> 
>>> def test_oddFrequencyCounter(oddFrequencyCounter):
...     paragraph = "The following is sample text I used to practice special characters using keybr.com: 112233445566778899 Saturn V rocket's first stage carries 203,400 gallons (770,000 liters) of kerosene fuel and 318,000 gallons (1.2 million liters) of liquid oxygen needed for combustion. At liftoff, the stage's five F-1 rocket engines ignite and produce 7.5 million pounds of thrust. To replace those goofy quantities with the far less retarded metric system (even though liters are considered part of the metric system they are the same as cubic deci-meters) you would say 770 cubic meters of kerosene {abbreviated as m3} and 1,204 m3 of liquid O2 [O2 is the symbol for oxygen]. We would also say it produced 33,600,000 newtons of force <abbreviated as N>. To add scientific notation {a way of writing numbers that allows you to write only as many digits of specificity as you would like} you can write 7.7 * 10 ^ 2 m3 of kerosene 1.204 * 10 ^ 3 m3 of O2 and 3.3 * 10 ^ 7 newtons."
...     assert oddFrequencyCounter(paragraph) == 40
```
```
>>> points: 1
>>> 
>>> def test_oddFrequencyCounter(oddFrequencyCounter):
...     paragraph = 'Data Science Africa in collaboration with The Swiss Federal Department of Foreign Affairs (FDFA), ETH Zurich, EPFL, the European Laboratory for Learning and Intelligent Systems (ELLIS), the Swiss National Computing Centre (CSCS), and the LUMI consortium launched the International Computation and AI Network (ICAIN) at the World Economic Forum (WEF) 2024 in Davos. Its mission is to develop AI technologies that benefit society as a whole, as well as being accessible to all and sustainable, thereby helping to reduce global inequality.'
...     assert oddFrequencyCounter(paragraph) == 20
```


Tests for q3 (points: 4)
```
>>> points: 4
>>> theGen = odd_squares_sum(12)
>>> assert next(theGen) == 1
>>> assert next(theGen) == 9
>>> assert next(theGen) == 25
>>> assert next(theGen) == 49
>>> assert next(theGen) == 81
>>> assert next(theGen) == 121
>>> del theGen
```


Tests for q4 (points: 4)
```
>>> points: 2
>>> assert (oddSumList == np.array([1, 9, 25, 49, 81, 121, 169, 225, 289, 361])).all()
```
```
>>> points: 2
>>> assert isinstance(oddSumList, np.ndarray)
```


Tests for q5 (points: 2)
```
>>> points: 2
>>> assert mergedList == [(1, 1), (4, 9), (0, 25), (4, 49), (1, 81), (1, 121), (4, 169), (0, 225), (4, 289), (1, 361)]
```


Tests for q6 (points: 5)
```
>>> points: 1
>>> np.testing.assert_equal(greatest_common_divisor(10, 15), 5)
```
```
>>> points: 1
>>> np.testing.assert_equal(greatest_common_divisor(15, 19), 1)
```
```
>>> points: 1
>>> np.testing.assert_equal(greatest_common_divisor(100, 105), 5)
```


Tests for q7 (points: 8)
```
>>> points: 2
>>> np.testing.assert_equal(get_3_nearest((3, 8), [(9, 3), (8, 5), (7, 6)]), [(7, 6), (8, 5), (9, 3)])
```
```
>>> points: 2
>>> np.testing.assert_equal(get_3_nearest((12, 8), [(5, 9), (9, 1), (2, 4), (13, 9), (10, 12)]), [(13, 9), (10, 12), (5, 9)])
```
```
>>> points: 2
>>> np.testing.assert_equal(get_3_nearest((5, 8), [(5, 9), (9, 1), (2, 4), (13, 9), (10, 12)]), [(5, 9), (2, 4), (13, 9)])
```


Tests for q8 (points: 4)
```
>>> points: 2
>>> np.testing.assert_array_equal(diagonal_vector([[1, 2, 3, 4, 5, 6, 7, 8], [9, 10, 11, 12, 13, 14, 15, 16], [17, 18, 19, 20, 21, 22, 23, 24], [25, 26, 27, 28, 29, 30, 31, 32], [33, 34, 35, 36, 37, 38, 39, 40], [41, 42, 43, 44, 45, 46, 47, 48], [49, 50, 51, 52, 53, 54, 55, 56], [57, 58, 59, 60, 61, 62, 63, 64]]), [1, 10, 19, 28, 37, 46, 55, 64])
```
```
>>> points: 2
>>> assert isinstance(diagonal_vector([[1, 2], [3, 4]]), np.ndarray)
```


Tests for q9 (points: 3)
```
>>> points: 1
>>> np.testing.assert_equal(flatten_reverse_lists([[2, 13, 44], [6, 7]]), [44, 13, 7, 6, 2])
```
```
>>> points: 1
>>> np.testing.assert_equal(flatten_reverse_lists([[2], [61, 34, 5, 8, 9]]), [61, 34, 9, 8, 5, 2])
```
